# Week 4 Exercises

See: _McKinney 3.1_ and [Python Documentation](https://docs.python.org/3/tutorial/datastructures.html) section 5 on data structures.

Modules 13 through 15 have all been about different built-in data types that can be used to create larger and more complex data structures.  We learned about tuples, lists, dictionaries, and how to access and process that data in new ways.

**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 20.1 Create a sorted list of names

Take the list of names below (some are first and last, others are first, middle, and last) and produce a final list of full names.

```
sort_names([
  ['Paul', 'Boal'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']])

['Paul Boal', 'Donald Duck', 'Kermit the Frog']
```

**Tip:** If you have a list of lists and want to sort it, Python can do that. It will simply sort using the first item in each of the lists inside your list.

```python
names = [['Paul', 'Boal'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']]
```

```python
def name_sorter(name_list):
    # Do something
    # Sort the list of names based on the first element in each sublist
    sorted_list = sorted(name_list, key=lambda x: x[-1])

    full_names = [' '.join(name) for name in sorted_list]

    return full_names
```

```python
names = [['Paul', 'Boal'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']]

assert name_sorter(names) == ["Paul Boal", "Donald Duck", "Kermit the Frog"]
```

### 20.2 Convert a dictionary to a list

In this exercise, we're going to take a dictionary that has families and their names, and convert that into a list of names.  The dictionary is structured such that the key is the family name, and the value for that key is a list of first names.

```
{
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}

[['Boal', 'Paul'],
 ['Boal', 'Anny'],
 ['Boal', 'Jim'],
 ['Lester-Boal', 'Elliie'],
 ['Lester-Boal', 'Ada'],
 ...
```

```python

```

```python
def dict_to_list(families):
   def name_sorter(name_list):
    # sort each sublist in the name_list
    sorted_names = [sorted(sublist) for sublist in name_list]

    # join the sorted names into full names as a string
    full_names = ["".join(sublist) for sublist in sorted_names]

    return full_names
```

```python
families = {
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}

dict_to_list({'Duck': ['Hewey','Dewey','Louie'],
              'Mouse': ['Mickey','Minney']})
```

### 20.3 Convert a list to a dictionary

Take the list of lists below and convert it into a dictionary where the entries in the dictionary are keyed off of whatever values are in column 0, and the dictionary values are the list of all the values that appear with that key from the input list.

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def list_to_dict():
    names = [['Boal', 'Paul'],
             ['Duck', 'Donald'],
             ['Duck', 'Daisy'],
             ['Boal', 'Ada'],
             ['Boal', 'Teddy'],
             ['Westhus', 'Eric']]

    families = {}

    # Iterate through the names list
    for name in names:
        # Extract the key (column 0) and value (column 1)
        key = name[0]
        value = name[1]

        # Check if the key already exists in the dictionary
        if key in families:
            # If it does, append the value to the existing list
            families[key].append(value)
        else:
            # If it doesn't, create a new list with the value
            families[key] = [value]

    return families

# Call the function and print the result
result = list_to_dict()
print(result)
```

Output:
```
{'Boal': ['Paul', 'Ada', 'Teddy'], 'Duck': ['Donald', 'Daisy'], 'Westhus': ['Eric']}
```

```python
list_to_dict()
```

Output:
```
{'Boal': ['Paul', 'Ada', 'Teddy'],
 'Duck': ['Donald', 'Daisy'],
 'Westhus': ['Eric']}
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
answer = {
 'Boal':    ['Paul', 'Ada', 'Teddy'],
 'Duck':    ['Donald', 'Daisy'],
 'Westhus': ['Eric']
}

assert (list_to_dict() == answer)
```

### 20.4 Join using a dictionary

We have a list of patients, diagnosis, and length of stay.  We also have a dictionary that contains diagnosis and average length of stay.  Produce an output list that lists the patient and an indicator if the patient's stay was 'too long', 'too short', 'just right'

<!-- grade_id: cell-f3813e33810dc898, points: 0, grade: false, solution: true -->
```python
def patient_los(avg_los, patients):

    conclusions = []
    # Put your solution here
    for patient in patients:
        # patient_name = patient[0]
        # diagnosis = patient[1]
        # length_of_stay = patient[2]
        patient_name, diagnosis, actual_los = patient
        if diagnosis in avg_los:
            avg_length = avg_los[diagnosis]

            if actual_los< avg_length:
                indicator = [patient_name,'too short']
            elif actual_los > avg_length:
                indicator = [patient_name,'too long']
            else:
                indicator = [patient_name,'just right']
        else:
            indicator = [patient_name,'No average length found for diagnosis']

        conclusions.append(indicator)
    return conclusions
```

```python
avg_los = {
    "Hemolytic jaundice and perinatal jaundice" : 2,
    "Medical examination/evaluation" : 3.2,
    "Liveborn" : 3.2,
    "Trauma to perineum and vulva" : 2.1,
    "Normal pregnancy and/or delivery" : 2,
    "Umbilical cord complication" : 2.1,
    "Forceps delivery" : 2.2,
    "Administrative/social admission" : 4.2,
    "Prolonged pregnancy" : 2.4,
    "Other complications of pregnancy" : 2.5
}

patients = [
    ['Boal', 'Medical examination/evaluation', 1.1],
    ['Boal', 'Other complications of pregnancy', 3.3],
    ['Jones', 'Liveborn', 3.2],
    ['Ashbury', 'Forceps delivery', 2.0]
]

patient_los(avg_los, patients)
```

Output:
```
[['Boal', 'too short'],
 ['Boal', 'too long'],
 ['Jones', 'just right'],
 ['Ashbury', 'too short']]
```

<!-- grade_id: cell-2af0a520bb56ee21, points: 10, grade: true, solution: false -->
```python
conclusions = [['Boal', 'too short'],
    ['Boal', 'too long'],
    ['Jones', 'just right'],
    ['Ashbury', 'too short']]

assert(patient_los(avg_los, patients) == conclusions)
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week04/week04_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

